## Exercise 05.1 (random numbers)

- Using the '`randint`' function from the '`random`' module (https://docs.python.org/3/library/random.html#random.randint), 
  develop a function `dice_roll` that emulates the roll of a dice with $n$ sides.

- For $n=6$, devise and implement a test to check that it is a fair dice.

#### (a) Dice roll code:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
for n in range(1, 20):
    for j in range(100):
        value = dice_roll(n) 
        assert value >= 1 and value <= n
```

#### (b) Test for fairness

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```