## Exercise 12.1

Create a class to represent vectors of arbitrary length and which is initialised with a list of values, e.g.:
```python
x = MyVector([0, 2, 4])
```

Equip the class with methods that:

1. Return the length of the vector (use name `size`)
2. Compute the norm of the vector $\sqrt{x \cdot x}$ (use name `norm`)
3. Compute the dot product of the vector with another vector (use name `dot`)

Test your implementation using two vectors of length 3. To help you get started, a skeleton of the class is provided below. Don't forget to use `self` where necessary.

```python
class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """
    
    def __init__(self, x):
        self.x = x
        
    # Return length of vector
    def size(self):
        # Add your code here
        pass  # This can be removed once the body is added
    
    # This allows access by index, e.g. y[2]
    def __getitem__(self, index):
        return self.x[index]

    # Return norm of vector
    def norm(self):
        # Add your code here
        pass  # This can be removed once the body is added
    
    # Return dot product of vector with another vector
    def dot(self, other):
        # Add your code here
        pass  # This can be removed once the body is added
```

### Solution

<!-- grade_id: cell-ff4d6cba19bb1b8e, points: 0, grade: false, solution: true -->
```python
import math  # Use this to get the sqrt

class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product 
    with another vector  """

    def __init__(self, x):
        self.x = x

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b601b5cc8cf8ab44, points: 0, grade: true, solution: false -->
```python
# Create two vectors
u = MyVector([1, 1, 2])
v = MyVector([2, 1, 1])

assert u.size() == 3
assert round(u.norm() - 2.449489742783178) == 0.0
assert round(u.dot(v) - 5.0, 10) == 0.0 
```